## Problem 2: Points to map (7 points)
 
The problem 2 this week continues the process that we started last week in Exercise 1, i.e. creating geometric point -objects and putting them into a map. 

**In this problem, our aim is to plot a map based on a set of x and y coordinates that are stored in a csv file.** 
The coordinates are in WGS84 decimal degrees, and the data is stored in `some_posts.csv` comma separated file in the folder `data`. First rows of the data look like this:
 
```
lat,lon,timestamp,userid
-24.980792492,31.484633302,2015-07-07 03:02,66487960
-25.499224667,31.508905612,2015-07-07 03:18,65281761
-24.342578456,30.930866066,2015-03-07 03:38,90916112
-24.85461393,31.519718439,2015-10-07 05:04,37959089
```

The data has 81379 rows and consists of locations and times of social media posts inside Kruger national park in South Africa:

| Column | Description |
|--------|-------------|
| lat | y-coordinate of the post |
| lon | x-coordinate of the post |
| timestamp | Time when the post was uploaded |
| userid | userid |

*Note: although the data is based on real social media data, it is heavily anonymized. Userids and timestamps have been randomized, i.e. they do not not match with real ones, also spatial accuracy of the data have been lowered.*

**First:**

- Import the needed modules
- Read the data from `data/some_posts.csv` into a Pandas dataframe called `data`
- Create an empty column called `geometry` where you will store shapely Point objects
- Instert Point objects into the column `geometry` based on the coordinate columns 

**HINTS:**

You want to create a Shapely point *on each row*, based on columns `'lon'` and `'lat'`. You can achieve this by using a `for`-loop and `.iterrows()` but this approach will be extremely slow! In stead, we recommend that you try one of these approaches:

- `zip` lon and lat columns and create the points using a for-loop and the zipped object, OR
- use the [apply method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.apply.html) to apply the shapely Point constructor on each row (only one line of code!)





<!-- grade_id: Create_points, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_create_points_test, points: 4, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Check the result
print("Number of rows:", len(data))
print(data['geometry'].head())
```

**Next:**
- Convert that DataFrame into a GeoDataFrame, [See hints](https://automating-gis-processes.github.io/2018/lessons/L2/exercise-2.html)

- Update the CRS for coordinate system as WGS84 (i.e. epsg code: 4326) **in the WKT format**

- Save the data into a Shapefile called `Kruger_posts.shp`



<!-- grade_id: Create_geodataframe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_create_geodataframe_test, points: 1.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Check the geodataframe head
print("Number of rows:", len(geo))
print(geo.head())
```

<!-- grade_id: problem_2_outputfile_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION

# Check that the output file exists
import os
assert os.path.isfile(fp), "output shapefile does not exist"
```

**Finally:** 
- Create a simple map of the points (one of these is enough):
    - Create the map using a the `.plot()` -funtion in Python. You might need to use the `matplotlib pyplot` `show()` method to display the image within this notebook. 
    - Create the map in a gis-software and upload the map as a png image to this repository (if you do this, remember to add a link to the image file to this notebook!)


<!-- grade_id: Plot_map, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_plot_map_check, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Well done! Now you can move on to [Exercise-2-problem-3.ipynb](Exercise-2-problem-3.ipynb)